---
title: 'R basics'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Exercise 1

```yaml
type: NormalExercise
key: 33cac5e5ab
xp: 100
```

type what ever you want in the editor. 
Use it as a calculator to do simple calculation.
For example, 3+9=? log(5)*sin(0.5*pi)

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: f326f9e1c7
xp: 100
```

`# complete the code below to get the output`
`mat<- matrix(c(3,1,5,7,3,5,0,2), )`

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
# complete the code below to get the output
mat<- matrix(c(3,1,5,7,3,5,0,2), )
```

`@solution`
```{r}
mat<- matrix(c(3,1,5,7,3,5,0,2), nrow=2, ncol=4, byrow=TRUE)
```

`@sct`
```{r}

```
